<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>9.22. Exercises &#8212; How to Think like a Computer Scientist: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.5.5" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=92D8BBF0"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.7/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="10. Lists" href="../Lists/toctree.html" />
    <link rel="prev" title="9.21. Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/sdmesa/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'https://sdmesa.github.io/thinkcspy/' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/sdmesa';
    eBookConfig.course = 'thinkcspy';
    eBookConfig.basecourse = 'thinkcspy';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.5.5';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/sdmesa/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          How to Think like a Computer Scientist
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/sdmesa/default/reportabug?course=thinkcspy&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">9.21. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">10. </span>Lists' data-toggle="tooltip" >              
    <a href="../Lists/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="exercises">
<h1><span class="section-number">9.22. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='q1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>What is the result of each of the following:</p>
<ol class="loweralpha simple">
<li><p>‘Python’[1]</p></li>
<li><p>“Strings are sequences of characters.”[5]</p></li>
<li><p>len(“wonderful”)</p></li>
<li><p>‘Mystery’[:4]</p></li>
<li><p>‘p’ in ‘Pineapple’</p></li>
<li><p>‘apple’ in ‘Pineapple’</p></li>
<li><p>‘pear’ not in ‘Pineapple’</p></li>
<li><p>‘apple’ &gt; ‘pineapple’</p></li>
<li><p>‘pineapple’ &lt; ‘Peach’</p></li>
</ol>
</div><div data-component="tab" data-tabname="Answer" >
<ol class="loweralpha simple">
<li><p>‘Python’[1] = ‘y’</p></li>
<li><p>‘Strings are sequences of characters.’[5] = ‘g’</p></li>
<li><p>len(‘wonderful’) = 9</p></li>
<li><p>‘Mystery’[:4] = ‘Myst’</p></li>
<li><p>‘p’ in ‘Pineapple’ = True</p></li>
<li><p>‘apple’ in ‘Pineapple’ = True</p></li>
<li><p>‘pear’ not in ‘Pineapple’ = True</p></li>
<li><p>‘apple’ &gt; ‘pineapple’ = False</p></li>
<li><p>‘pineapple’ &lt; ‘Peach’ = False</p></li>
</ol>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_dc2457710a924d9283b12f42a31d2b27_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course;
                    this.page.url = 'https://disqus_dc2457710a924d9283b12f42a31d2b27_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course;
            disqus_url = 'http://disqus_dc2457710a924d9283b12f42a31d2b27_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_dc2457710a924d9283b12f42a31d2b27" onclick="disqus_dc2457710a924d9283b12f42a31d2b27_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_dc2457710a924d9283b12f42a31d2b27']").attr('data-disqus-identifier', 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_2 >
<div id=ex_8_2_question class="ac_question col-md-12">
<p>In Robert McCloskey’s
book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack,
Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">prefixes</span> <span class="o">=</span> <span class="s2">&quot;JKLMNOPQ&quot;</span>
<span class="n">suffix</span> <span class="o">=</span> <span class="s2">&quot;ack&quot;</span>

<span class="k">for</span> <span class="n">p</span> <span class="ow">in</span> <span class="n">prefixes</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="n">p</span> <span class="o">+</span> <span class="n">suffix</span><span class="p">)</span>
</pre></div>
</div>
<p>Of course, that’s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>

</div>
<textarea data-lang="python" id="ex_8_2_editor" data-question_label="9.22.1" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

prefixes = "JKLMNOPQ"
suffix = "ack"

for p in prefixes:
    print(p + suffix)

# Fix the loop to get the correct output.

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        o = self.getOutput()
        code = self.getEditorText()
        self.assertIn("if", code, "Needs a conditional.")
        self.assertIn("for", code, "Needs a loop.")
        self.assertIn("Jack", o, "J + ack = Jack")
        self.assertIn("Kack", o, "K + ack = Kack")
        self.assertIn("Lack", o, "L + ack = Lack")
        self.assertIn("Mack", o, "M + ack = Mack")
        self.assertIn("Nack", o, "N + ack = Nack")
        self.assertIn("Ouack", o, "Don't forget the misspellings. Quack is required.")
        self.assertIn("Pack", o, "P + ack = Pack")
        self.assertIn("Quack", o, "Don't forget the misspellings. Quack is required.")
        self.assertNotIn("Oack", o, "Account for the misspellings. Qack should not be in output.")
        self.assertNotIn("Qack", o, "Account for the misspellings. Qack should not be in output.")
myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_3 >
<div id=ex_8_3_question class="ac_question col-md-12">
<p>Assign to a variable in your program a triple-quoted string that contains
your favorite paragraph of text - perhaps a poem, a speech, instructions
to bake a cake, some inspirational verses, etc.</p>
<p>Write a function that counts the number of alphabetic characters (a through z, or A through Z) in your text and then keeps track of how many are the letter ‘e’.  Your function should print an analysis of the text like this:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Your</span> <span class="n">text</span> <span class="n">contains</span> <span class="mi">243</span> <span class="n">alphabetic</span> <span class="n">characters</span><span class="p">,</span> <span class="n">of</span> <span class="n">which</span> <span class="mi">109</span> <span class="p">(</span><span class="mf">44.8</span><span class="o">%</span><span class="p">)</span> <span class="n">are</span> <span class="s1">&#39;e&#39;</span><span class="o">.</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_8_3_editor" data-question_label="9.22.2" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def count(p):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        string1 = "e"
        string2 = "eieio"
        string3 = "eeeeeeeeeeee"
        string4 = "elephant"
        self.assertEqual(count(string1), 1, "Just one e")
        self.assertEqual(count(""), 0, "Empty string")
        self.assertEqual(count(string2), 2, "Two")
        self.assertEqual(count(string3), 12, "Twelve")
        self.assertNotEqual(count(string4), 3, "Has two Es")
myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q3_answer >
<div id=str_q3_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q3_answer_editor" data-question_label="9.22.3" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def count(p):
    lows = "abcdefghijklmnopqrstuvwxyz"
    ups =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"

    numberOfe = 0
    totalChars = 0
    for achar in p:
        if achar in lows or achar in ups:
            totalChars = totalChars + 1
            if achar == 'e':
                numberOfe = numberOfe + 1

    percent_with_e = (numberOfe / totalChars) * 100
    print("Your text contains", totalChars, "alphabetic characters of which", numberOfe, "(", percent_with_e, "%)", "are 'e'.")


p = '''
"If the automobile had followed the same development cycle as the computer, a
Rolls-Royce would today cost $100, get a million miles per gallon, and explode
once a year, killing everyone inside."
-Robert Cringely
'''

count(p)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_bf88b1c4616d43f289c798b56a43b01c_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course;
                    this.page.url = 'https://disqus_bf88b1c4616d43f289c798b56a43b01c_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course;
            disqus_url = 'http://disqus_bf88b1c4616d43f289c798b56a43b01c_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_bf88b1c4616d43f289c798b56a43b01c" onclick="disqus_bf88b1c4616d43f289c798b56a43b01c_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_bf88b1c4616d43f289c798b56a43b01c']").attr('data-disqus-identifier', 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_4 >
<div id=ex_8_4_question class="ac_question col-md-12">
<p>Print out a neatly formatted multiplication table, up to 12 x 12.</p>

</div>
<textarea data-lang="python" id="ex_8_4_editor" data-question_label="9.22.4" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='q5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_7_10 >
<div id=ex_7_10_question class="ac_question col-md-12">
<p>Write a function that will return the number of digits in an integer.</p>

</div>
<textarea data-lang="python" id="ex_7_10_editor" data-question_label="9.22.5" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def numDigits(n):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
      self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
      self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
      self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q5_answer >
<div id=str_q5_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q5_answer_editor" data-question_label="9.22.6" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def numDigits(n):
    n_str = str(n)
    return len(n_str)


print(numDigits(50))
print(numDigits(20000))
print(numDigits(1))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
                    this.page.url = 'https://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_bfd6f74a183c4682b29c72c4411200fb" onclick="disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_bfd6f74a183c4682b29c72c4411200fb']").attr('data-disqus-identifier', 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_5 >
<div id=ex_8_5_question class="ac_question col-md-12">
<p>Write a function that reverses its string argument.</p>

</div>
<textarea data-lang="python" id="ex_8_5_editor" data-question_label="9.22.7" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def reverse(astring):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
      self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
      self.assertEqual(reverse(""),"","Tested reverse on input of ''")




myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Write a function that mirrors its string argument,
generating a string containing the original string and the string backwards.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_6 >
<div id=ex_8_6_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ex_8_6_editor" data-question_label="9.22.8" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def mirror(mystr):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(mirror("good"),"gooddoog","Tested mirror on input of 'good'")
        self.assertEqual(mirror("Python"),"PythonnohtyP","Tested mirror on input of 'Python'")
        self.assertEqual(mirror(""),"","Tested mirror on input of ''")
        self.assertEqual(mirror("a"),"aa","Tested mirror on input of 'a'")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q7_answer >
<div id=str_q7_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q7_answer_editor" data-question_label="9.22.9" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def reverse(mystr):
    reversed = ''
    for char in mystr:
        reversed = char + reversed
    return reversed

def mirror(mystr):
    return mystr + reverse(mystr)

testEqual(mirror('good'), 'gooddoog')
testEqual(mirror('Python'), 'PythonnohtyP')
testEqual(mirror(''), '')
testEqual(mirror('a'), 'aa')
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_70b7ac515456497c952a2de5caa27ab9_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
                    this.page.url = 'https://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
            disqus_url = 'http://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_70b7ac515456497c952a2de5caa27ab9" onclick="disqus_70b7ac515456497c952a2de5caa27ab9_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_70b7ac515456497c952a2de5caa27ab9']").attr('data-disqus-identifier', 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_7 >
<div id=ex_8_7_question class="ac_question col-md-12">
<p>Write a function that removes all occurrences of a given letter from a string.</p>

</div>
<textarea data-lang="python" id="ex_8_7_editor" data-question_label="9.22.10" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def remove_letter(theLetter, theString):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
      self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
      self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")



myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_8 >
<div id=ex_8_8_question class="ac_question col-md-12">
<p>Write a function that recognizes palindromes. (Hint: use your <code class="docutils literal notranslate"><span class="pre">reverse</span></code> function to make this easy!).</p>

</div>
<textarea data-lang="python" id="ex_8_8_editor" data-question_label="9.22.11" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def is_palindrome(myStr):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_palindrome("abba"),True,"Tested is_palindrome on input of 'abba'")
        self.assertEqual(is_palindrome("abab"),False,"Tested is_palindrome on input of 'abab'")
        self.assertEqual(is_palindrome("straw warts"),True,"Tested is_palindrome on input of 'straw warts'")
        self.assertEqual(is_palindrome("a"),True,"Tested is_palindrome on input of 'a'")
        self.assertEqual(is_palindrome(""),True,"Tested is_palindrome on input of ''")




myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q9_answer >
<div id=str_q9_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q9_answer_editor" data-question_label="9.22.12" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def reverse(mystr):
    reversed = ''
    for char in mystr:
        reversed = char + reversed
    return reversed

def is_palindrome(myStr):
    if myStr in reverse(myStr):
        return True
    else:
        return False

testEqual(is_palindrome('abba'), True)
testEqual(is_palindrome('abab'), False)
testEqual(is_palindrome('straw warts'), True)
testEqual(is_palindrome('a'), True)
testEqual(is_palindrome(''), True)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_605923545bb849f7b8d41bbf823518e9_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course;
                    this.page.url = 'https://disqus_605923545bb849f7b8d41bbf823518e9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course;
            disqus_url = 'http://disqus_605923545bb849f7b8d41bbf823518e9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_605923545bb849f7b8d41bbf823518e9" onclick="disqus_605923545bb849f7b8d41bbf823518e9_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_605923545bb849f7b8d41bbf823518e9']").attr('data-disqus-identifier', 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p>Write a function that counts how many non-overlapping occurences of a substring appear in a string.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_9 >
<div id=ex_8_9_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ex_8_9_editor" data-question_label="9.22.13" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def count(substr,theStr):
    # your code here


====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(count("is","Mississippi"),2,"Tested count on inputs of 'is' and 'Mississippi'")
          self.assertEqual(count("an","banana"),2,"Tested count on inputs of 'an' and 'banana'")
          self.assertEqual(count("ana","banana"),1,"Tested count on inputs of 'ana' and 'banana'")
          self.assertEqual(count("nana","banana"),1,"Tested count on inputs of 'nana' and 'banana'")
          self.assertEqual(count("nanan","banana"),0,"Tested count on inputs of 'nanan' and 'banana'")
          self.assertEqual(count("aaa","aaaaaa"),2,"Tested count on input of 'aaa' and 'aaaaaa'")




myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q11' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_10 >
<div id=ex_8_10_question class="ac_question col-md-12">
<p>Write a function that removes the first occurrence of a string from another string.</p>

</div>
<textarea data-lang="python" id="ex_8_10_editor" data-question_label="9.22.14" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def remove(substr,theStr):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(remove("an","banana"),"bana","Tested remove on inputs of 'an' and 'banana'")
            self.assertEqual(remove("cyc","bicycle"),"bile","Tested remove on inputs of 'cyc' and 'bicycle'")
            self.assertEqual(remove("iss","Mississippi"),"Missippi","Tested remove on inputs of 'iss' and 'Mississippi'")
            self.assertEqual(remove("egg","bicycle"),"bicycle","Tested remove on inputs of 'egg' and 'bicycle'")



myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q11_answer >
<div id=str_q11_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q11_answer_editor" data-question_label="9.22.15" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def remove(substr,theStr):
    index = theStr.find(substr)
    if index < 0: # substr doesn't exist in theStr
        return theStr
    return_str = theStr[:index] + theStr[index+len(substr):]
    return return_str

testEqual(remove('an', 'banana'), 'bana')
testEqual(remove('cyc', 'bicycle'), 'bile')
testEqual(remove('iss', 'Mississippi'), 'Missippi')
testEqual(remove('egg', 'bicycle'), 'bicycle')
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_2f2772134b604a6498748138542d312d_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course;
                    this.page.url = 'https://disqus_2f2772134b604a6498748138542d312d_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course;
            disqus_url = 'http://disqus_2f2772134b604a6498748138542d312d_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_2f2772134b604a6498748138542d312d" onclick="disqus_2f2772134b604a6498748138542d312d_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_2f2772134b604a6498748138542d312d']").attr('data-disqus-identifier', 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_11 >
<div id=ex_8_11_question class="ac_question col-md-12">
<p>Write a function that removes all occurrences of a string from another string.</p>

</div>
<textarea data-lang="python" id="ex_8_11_editor" data-question_label="9.22.16" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def remove_all(substr,theStr):
    # your code here



====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_all("an","banana"),"ba","Tested remove_all on inputs of 'an' and 'banana'")
      self.assertEqual(remove_all("cyc","bicycle"),"bile","Tested remove_all on inputs of 'cyc' and 'bicycle'")
      self.assertEqual(remove_all("iss","Mississippi"),"Mippi","Tested remove_all on inputs of 'iss' and 'Mississippi'")
      self.assertEqual(remove_all("eggs","bicycle"),"bicycle","Tested remove_all on inputs of 'eggs' and 'bicycle'")



myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q13' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_12 >
<div id=ex_8_12_question class="ac_question col-md-12">
<p>Here is another interesting L-System called a Hilbert curve.  Use 90 degrees:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">L</span>
<span class="n">L</span> <span class="o">-&gt;</span> <span class="o">+</span><span class="n">RF</span><span class="o">-</span><span class="n">LFL</span><span class="o">-</span><span class="n">FR</span><span class="o">+</span>
<span class="n">R</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">LF</span><span class="o">+</span><span class="n">RFR</span><span class="o">+</span><span class="n">FL</span><span class="o">-</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_8_12_editor" data-question_label="9.22.17" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q13_answer >
<div id=str_q13_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q13_answer_editor" data-question_label="9.22.18" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'L':
        newstr = '+RF-LFL-FR+'   # Rule 1
    elif ch == 'R':
        newstr = '-LF+RFR+FL-'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(4, "L")  # create the string
    print(inst)
    t = turtle.Turtle()           # create the turtle
    wn = turtle.Screen()

    t.up()
    t.back(200)
    t.down()
    t.speed(9)
    drawLsystem(t, inst, 90, 5)   # draw the picture
                                  # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_ab823200fac64461a9e88f53b75f5795_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course;
                    this.page.url = 'https://disqus_ab823200fac64461a9e88f53b75f5795_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course;
            disqus_url = 'http://disqus_ab823200fac64461a9e88f53b75f5795_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_ab823200fac64461a9e88f53b75f5795" onclick="disqus_ab823200fac64461a9e88f53b75f5795_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_ab823200fac64461a9e88f53b75f5795']").attr('data-disqus-identifier', 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_13 >
<div id=ex_8_13_question class="ac_question col-md-12">
<p>Here is a dragon curve.  Use 90 degrees.:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">FX</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">X</span><span class="o">+</span><span class="n">YF</span><span class="o">+</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">FX</span><span class="o">-</span><span class="n">Y</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_8_13_editor" data-question_label="9.22.19" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='q15' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_14 >
<div id=ex_8_14_question class="ac_question col-md-12">
<p>Here is something called an arrowhead curve.  Use 60 degrees.:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">YF</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">YF</span><span class="o">+</span><span class="n">XF</span><span class="o">+</span><span class="n">Y</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="n">XF</span><span class="o">-</span><span class="n">YF</span><span class="o">-</span><span class="n">X</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_8_14_editor" data-question_label="9.22.20" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q15_answer >
<div id=str_q15_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q15_answer_editor" data-question_label="9.22.21" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'X':
        newstr = 'YF+XF+Y'   # Rule 1
    elif ch == 'Y':
        newstr = 'XF-YF-X'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(5, "YF")  # create the string
    print(inst)
    t = turtle.Turtle()            # create the turtle
    wn = turtle.Screen()

    t.speed(9)
    drawLsystem(t, inst, 60, 5)    # draw the picture
                                   # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_9b2dfba083a64d5c894f873af2e93a1b_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course;
                    this.page.url = 'https://disqus_9b2dfba083a64d5c894f873af2e93a1b_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course;
            disqus_url = 'http://disqus_9b2dfba083a64d5c894f873af2e93a1b_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_9b2dfba083a64d5c894f873af2e93a1b" onclick="disqus_9b2dfba083a64d5c894f873af2e93a1b_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_9b2dfba083a64d5c894f873af2e93a1b']").attr('data-disqus-identifier', 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_15 >
<div id=ex_8_15_question class="ac_question col-md-12">
<p>Try the Peano-Gosper curve.  Use 60 degrees.:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">FX</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">X</span><span class="o">+</span><span class="n">YF</span><span class="o">++</span><span class="n">YF</span><span class="o">-</span><span class="n">FX</span><span class="o">--</span><span class="n">FXFX</span><span class="o">-</span><span class="n">YF</span><span class="o">+</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">FX</span><span class="o">+</span><span class="n">YFYF</span><span class="o">++</span><span class="n">YF</span><span class="o">+</span><span class="n">FX</span><span class="o">--</span><span class="n">FX</span><span class="o">-</span><span class="n">Y</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_8_15_editor" data-question_label="9.22.22" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='q17' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_16 >
<div id=ex_8_16_question class="ac_question col-md-12">
<blockquote>
<div><p>The Sierpinski Triangle.  Use 60 degrees.:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">FXF</span><span class="o">--</span><span class="n">FF</span><span class="o">--</span><span class="n">FF</span>
<span class="n">F</span> <span class="o">-&gt;</span> <span class="n">FF</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="o">--</span><span class="n">FXF</span><span class="o">++</span><span class="n">FXF</span><span class="o">++</span><span class="n">FXF</span><span class="o">--</span>
</pre></div>
</div>
</div></blockquote>

</div>
<textarea data-lang="python" id="ex_8_16_editor" data-question_label="9.22.23" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q17_answer >
<div id=str_q17_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q17_answer_editor" data-question_label="9.22.24" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'F':
        newstr = 'FF'   # Rule 1
    elif ch == 'X':
        newstr = '--FXF++FXF++FXF--'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(5, "FXF--FF--FF")   # create the string
    print(inst)
    t = turtle.Turtle()           # create the turtle
    wn = turtle.Screen()
    t.up()
    t.back(200)
    t.right(90)
    t.forward(100)
    t.left(90)
    t.down()
    t.speed(9)

    drawLsystem(t, inst, 60, 5)   # draw the picture
                                  # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_912a5f19d3964dc2af7a067dcd832c7e_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course;
                    this.page.url = 'https://disqus_912a5f19d3964dc2af7a067dcd832c7e_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course;
            disqus_url = 'http://disqus_912a5f19d3964dc2af7a067dcd832c7e_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_912a5f19d3964dc2af7a067dcd832c7e" onclick="disqus_912a5f19d3964dc2af7a067dcd832c7e_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_912a5f19d3964dc2af7a067dcd832c7e']").attr('data-disqus-identifier', 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_17 >
<div id=ex_8_17_question class="ac_question col-md-12">
<p>Write a function that implements a substitution cipher.  In a substitution
cipher one letter is substituted for another to garble the message.  For
example A -&gt; Q, B -&gt; T, C -&gt; G etc.  your function should take two
parameters, the message you want to encrypt, and a string that represents
the mapping of the 26 letters in the alphabet.  Your function should
return a string that is the encrypted version of the message.</p>

</div>
<textarea data-lang="python" id="ex_8_17_editor" data-question_label="9.22.25" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</li>
<li><div id='q19' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_18 >
<div id=ex_8_18_question class="ac_question col-md-12">
<p>Write a function that decrypts the message from the previous exercise.  It
should also take two parameters.  The encrypted message,
and the mixed up alphabet.  The function should return a string that is
the same as the original unencrypted message.</p>

</div>
<textarea data-lang="python" id="ex_8_18_editor" data-question_label="9.22.26" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q19_answer >
<div id=str_q19_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q19_answer_editor" data-question_label="9.22.27" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def encrypt(message, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    encrypted = ''
    for char in message:
        if char == ' ':
            encrypted = encrypted + ' '
        else:
            pos = alphabet.index(char)
            encrypted = encrypted + cipher[pos]
    return encrypted

def decrypt(encrypted, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    decrypted = ''
    for char in encrypted:
        if char == ' ':
            decrypted = decrypted + ' '
        else:
            pos = cipher.index(char)
            decrypted = decrypted + alphabet[pos]
    return decrypted


cipher = "badcfehgjilknmporqtsvuxwzy"

encrypted = encrypt('hello world', cipher)
print encrypted

decrypted = decrypt(encrypted, cipher)
print(decrypted)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_d7d1ca00bfff4e7bacf886386fb3302e_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course;
                    this.page.url = 'https://disqus_d7d1ca00bfff4e7bacf886386fb3302e_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course;
            disqus_url = 'http://disqus_d7d1ca00bfff4e7bacf886386fb3302e_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_d7d1ca00bfff4e7bacf886386fb3302e" onclick="disqus_d7d1ca00bfff4e7bacf886386fb3302e_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_d7d1ca00bfff4e7bacf886386fb3302e']").attr('data-disqus-identifier', 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_19 >
<div id=ex_8_19_question class="ac_question col-md-12">
<p>Write a function called  <code class="docutils literal notranslate"><span class="pre">remove_dups</span></code> that takes a string and creates a new string by only adding those characters that are not already present.  In other words,
there will never be a duplicate letter added to the new string.</p>

</div>
<textarea data-lang="python" id="ex_8_19_editor" data-question_label="9.22.28" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def remove_dups(astring):
    # your code here


print(remove_dups("mississippi"))   #should print misp

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_dups("pooh"),"poh","Tested remove_dups on string 'pooh'")
      self.assertEqual(remove_dups("mississippi"),"misp","Tested remove_dups on string 'mississippi'")
      self.assertEqual(remove_dups("potato"),"pota","Tested remove_dups on string 'potato'")
      self.assertEqual(remove_dups("bookkeeper"),"bokepr","Tested remove_dups on string 'bookkeeper'")
      self.assertEqual(remove_dups("oo"),"o","Tested remove_dups on string 'oo'")

myTests().main()
</textarea>
</div>
</div>
</li>
<li><div id='q21' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_20 >
<div id=ex_8_20_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">rot13</span></code> that uses the Caesar cipher to encrypt a message.
The Caesar cipher works like a substitution cipher but each character is replaced
by the character 13 characters to ‘its right’ in the alphabet.  So for example
the letter a becomes the letter n.  If a letter is past the middle of the alphabet
then the counting wraps around to the letter a again, so n becomes a, o becomes b
and so on.  <em>Hint:</em> Whenever you talk about things wrapping around its a good idea
to think of modulo arithmetic.</p>

</div>
<textarea data-lang="python" id="ex_8_20_editor" data-question_label="9.22.29" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def rot13(mess):
    # Your code here

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('Since rot13 is symmetric you should see this message')))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=str_q21_answer >
<div id=str_q21_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="str_q21_answer_editor" data-question_label="9.22.30" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def rot13(mess):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted = ''
    for char in mess:
        if char == ' ':
            encrypted = encrypted + ' '
        else:
            rotated_index = alphabet.index(char) + 13
            if rotated_index < 26:
                encrypted = encrypted + alphabet[rotated_index]
            else:
                encrypted = encrypted + alphabet[rotated_index % 26]
    return encrypted

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('since rot thirteen is symmetric you should see this message')))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_49e1151bb7864a3287a6b6ae1c84db16_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course;
                    this.page.url = 'https://disqus_49e1151bb7864a3287a6b6ae1c84db16_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course;
            disqus_url = 'http://disqus_49e1151bb7864a3287a6b6ae1c84db16_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_49e1151bb7864a3287a6b6ae1c84db16" onclick="disqus_49e1151bb7864a3287a6b6ae1c84db16_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_49e1151bb7864a3287a6b6ae1c84db16']").attr('data-disqus-identifier', 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_8_22 >
<div id=ex_8_22_question class="ac_question col-md-12">
<p>Modify this code so it prints each subtotal, the total cost, and average price
to exactly two decimal places.</p>

</div>
<textarea data-lang="python" id="ex_8_22_editor" data-question_label="9.22.31" 
    data-hidecode="true"  data-timelimit=60 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def checkout():
    total = 0
    count = 0
    moreItems = True
    while moreItems:
        price = float(input('Enter price of item (0 when done): '))
        if price != 0:
            count = count + 1
            total = total + price
            print('Subtotal: $', total)
        else:
            moreItems = False
    average = total / count
    print('Total items:', count)
    print('Total $', total)
    print('Average price per item: $', average)

checkout()
</textarea>
</div>
</div>
</li>
</ol>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 9.21. Glossary' data-toggle="tooltip">
    <a href="Glossary.html" aria-label="9.21. Glossary">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 10. Lists' data-toggle="tooltip" >
    <a href="../Lists/toctree.html" aria-label="10. Lists">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2014 Brad Miller, David Ranum, Created using Runestone Interactive.
      Last updated on Nov 14, 2021.<br/>
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.5.5.
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>