<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>16.8. Programming Exercises &#8212; How to Think like a Computer Scientist: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.5.5" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=92D8BBF0"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.7/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="16.9. Exercises" href="Exercises.html" />
    <link rel="prev" title="16.7. Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/sdmesa/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'https://sdmesa.github.io/thinkcspy/' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/sdmesa';
    eBookConfig.course = 'thinkcspy';
    eBookConfig.basecourse = 'thinkcspy';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.5.5';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/mesa_logo.png") no-repeat 0px 0px;' href="/sdmesa/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          How to Think like a Computer Scientist
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- WWW external links dropdown -->

        <!--
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
        </li>-->

        <!-- end WWW external links dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/sdmesa/default/reportabug?course=thinkcspy&page=ProgrammingExercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">16.7. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">16.9. </span>Exercises' data-toggle="tooltip" >              
    <a href="Exercises.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="programming-exercises">
<h1><span class="section-number">16.8. </span>Programming Exercises<a class="headerlink" href="#programming-exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" data-question_label="16.8.1" id="recursion_ex_1" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_1 >
<div id=ex_rec_1_question class="ac_question col-md-12">
<p>Write a recursive function to compute the factorial of a number.</p>

</div>
<textarea data-lang="python" id="ex_rec_1_editor" data-question_label="16.8.2" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def computeFactorial(number):
    #your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(computeFactorial(0),1,"Tested computeFactorial on input 0")
        self.assertEqual(computeFactorial(1),1,"Tested computeFactorail on input 1")
        self.assertEqual(computeFactorial(2),2,"Tested computeFactorial on input 2")
        self.assertEqual(computeFactorial(3),6,"Tested computeFactorial on input 3")
        self.assertEqual(computeFactorial(4),24,"Tested computeFactorial on input 4")
        self.assertEqual(computeFactorial(8),40320,"Tested computeFactorial on input 8")
        self.assertEqual(computeFactorial(-5),None,"Tested computeFactorial on a negative number - make sure to handle this case")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.3" id="recursion_ex_2" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_2 >
<div id=ex_rec_2_question class="ac_question col-md-12">
<p>Write a recursive function to reverse a list.</p>

</div>
<textarea data-lang="python" id="ex_rec_2_editor" data-question_label="16.8.4" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def reverseList(lst):
    #your code here


====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(reverseList([1,2,3,4,5]), [5,4,3,2,1], "Your function failed with input [1,2,3,4,5]")
        self.assertEqual(reverseList(['Hello','World','!']), ['!','World','Hello'], "Your function failed with input ['Hello,'World','!']")
        self.assertEqual(reverseList(['Python',100,'35','Computer Science']), ['Computer Science', '35', 100, 'Python'], "Your function failed with input ['Python,100,'35','Computer Science']")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.5" id="recursion_ex_3" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_3 >
<div id=ex_rec_3_question class="ac_question col-md-12">
<p>Modify the recursive tree program using one or all of the following
ideas:</p>
<ul class="simple">
<li><p>Modify the thickness of the branches so that as the <code class="docutils literal notranslate"><span class="pre">branchLen</span></code>
gets smaller, the line gets thinner.</p></li>
<li><p>Modify the color of the branches so that as the <code class="docutils literal notranslate"><span class="pre">branchLen</span></code> gets
very short it is colored like a leaf.</p></li>
<li><p>Modify the angle used in turning the turtle so that at each branch
point the angle is selected at random in some range. For example
choose the angle between 15 and 45 degrees. Play around to see
what looks good.</p></li>
<li><p>Modify the <code class="docutils literal notranslate"><span class="pre">branchLen</span></code> recursively so that instead of always
subtracting the same amount you subtract a random amount in some
range.</p></li>
</ul>
<p>If you implement all of the above ideas you will have a very
realistic looking tree.</p>

</div>
<textarea data-lang="python" id="ex_rec_3_editor" data-question_label="16.8.6" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.7" id="recursion_ex_4" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_4 >
<div id=ex_rec_4_question class="ac_question col-md-12">
<p>Find or invent an algorithm for drawing a fractal mountain. Hint: One
approach to this uses triangles again.</p>

</div>
<textarea data-lang="python" id="ex_rec_4_editor" data-question_label="16.8.8" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.9" id="recursion_ex_5" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_5 >
<div id=ex_rec_5_question class="ac_question col-md-12">
<p>Write a recursive function to compute the Fibonacci sequence. How
does the performance of the recursive function compare to that of an
iterative version?</p>

</div>
<textarea data-lang="python" id="ex_rec_5_editor" data-question_label="16.8.10" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.11" id="recursion_ex_6" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_6 >
<div id=ex_rec_6_question class="ac_question col-md-12">
<p>Implement a solution to the Tower of Hanoi using three stacks to keep
track of the disks.</p>

</div>
<textarea data-lang="python" id="ex_rec_6_editor" data-question_label="16.8.12" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.13" id="recursion_ex_7" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_7 >
<div id=ex_rec_7_question class="ac_question col-md-12">
<p>Using the turtle graphics module, write a recursive program to
display a Hilbert curve.</p>

</div>
<textarea data-lang="python" id="ex_rec_7_editor" data-question_label="16.8.14" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.15" id="recursion_ex_8" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_8 >
<div id=ex_rec_8_question class="ac_question col-md-12">
<p>Using the turtle graphics module, write a recursive program to
display a Koch snowflake.</p>

</div>
<textarea data-lang="python" id="ex_rec_8_editor" data-question_label="16.8.16" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.17" id="recursion_ex_9" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_9 >
<div id=ex_rec_9_question class="ac_question col-md-12">
<p>Write a program to solve the following problem: You have two jugs: a
4-gallon jug and a 3-gallon jug. Neither of the jugs have markings on
them. There is a pump that can be used to fill the jugs with water.
How can you get exactly two gallons of water in the 4-gallon jug?</p>

</div>
<textarea data-lang="python" id="ex_rec_9_editor" data-question_label="16.8.18" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.19" id="recursion_ex_10" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_10 >
<div id=ex_rec_10_question class="ac_question col-md-12">
<p>Generalize the problem above so that the parameters to your solution
include the sizes of each jug and the final amount of water to be
left in the larger jug.</p>

</div>
<textarea data-lang="python" id="ex_rec_10_editor" data-question_label="16.8.20" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.21" id="recursion_ex_11" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_11 >
<div id=ex_rec_11_question class="ac_question col-md-12">
<p>Write a program that solves the following problem: Three missionaries
and three cannibals come to a river and find a boat that holds two
people. Everyone must get across the river to continue on the
journey. However, if the cannibals ever outnumber the missionaries on
either bank, the missionaries will be eaten. Find a series of
crossings that will get everyone safely to the other side of the
river.</p>

</div>
<textarea data-lang="python" id="ex_rec_11_editor" data-question_label="16.8.22" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.23" id="recursion_ex_12" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_12 >
<div id=ex_rec_12_question class="ac_question col-md-12">
<p>Modify the Tower of Hanoi program using turtle graphics to animate
the movement of the disks. Hint: You can make multiple turtles and
have them shaped like rectangles.</p>

</div>
<textarea data-lang="python" id="ex_rec_12_editor" data-question_label="16.8.24" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="16.8.25" id="recursion_ex_13" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_rec_13 >
<div id=ex_rec_13_question class="ac_question col-md-12">
<p>Pascal’s triangle is a number triangle with numbers arranged in
staggered rows such that</p>
<div class="math notranslate nohighlight">
\[a_{nr} = {n! \over{r! (n-r)!}}\]</div>
<p>This equation is the equation for a binomial coefficient. You can
build Pascal’s triangle by adding the two numbers that are diagonally
above a number in the triangle. An example of Pascal’s triangle is
shown below.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>        <span class="mi">1</span>
      <span class="mi">1</span>   <span class="mi">1</span>
    <span class="mi">1</span>   <span class="mi">2</span>   <span class="mi">1</span>
  <span class="mi">1</span>   <span class="mi">3</span>   <span class="mi">3</span>   <span class="mi">1</span>
<span class="mi">1</span>   <span class="mi">4</span>   <span class="mi">6</span>   <span class="mi">4</span>   <span class="mi">1</span>
</pre></div>
</div>
<p>Write a program that prints out Pascal’s triangle. Your program
should accept a parameter that tells how many rows of the triangle to
print.</p>

</div>
<textarea data-lang="python" id="ex_rec_13_editor" data-question_label="16.8.26" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    </div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 16.7. Glossary' data-toggle="tooltip">
    <a href="Glossary.html" aria-label="16.7. Glossary">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 16.9. Exercises' data-toggle="tooltip" >
    <a href="Exercises.html" aria-label="16.9. Exercises">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2014 Brad Miller, David Ranum, and Dave Parillo.
      Last updated on Nov 14, 2021.<br/>
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.5.5.
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>