<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>7.10. Exercises &#8212; How to Think like a Computer Scientist: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.5.5" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=92D8BBF0"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.7/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="8. More About Iteration" href="../MoreAboutIteration/toctree.html" />
    <link rel="prev" title="7.9. Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/sdmesa/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'https://sdmesa.github.io/thinkcspy/' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/sdmesa';
    eBookConfig.course = 'thinkcspy';
    eBookConfig.basecourse = 'thinkcspy';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.5.5';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/sdmesa/default/user/login">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          How to Think like a Computer Scientist
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- WWW external links dropdown -->

        <!--
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
        </li>-->

        <!-- end WWW external links dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/sdmesa/default/reportabug?course=thinkcspy&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">7.9. </span>Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">8. </span>More About Iteration' data-toggle="tooltip" >              
    <a href="../MoreAboutIteration/toctree.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="exercises">
<h1><span class="section-number">7.10. </span>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" data-question_label="7.10.1" id="selection_ex_1" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='q1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_1 >
<div id=ex_6_1_question class="ac_question col-md-12">
<p>What do these expressions evaluate to?</p>
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">==</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">!=</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">&gt;=</span> <span class="pre">4</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">(3</span> <span class="pre">&lt;</span> <span class="pre">4)</span></code></p></li>
</ol>

</div>
<textarea data-lang="python" id="ex_6_1_editor" data-question_label="7.10.2" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<ol class="arabic simple">
<li><p>True</p></li>
<li><p>False</p></li>
<li><p>False</p></li>
<li><p>False</p></li>
</ol>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_eb4a097382404ffe81300aac5744e3fe_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_eb4a097382404ffe81300aac5744e3fe_' + eBookConfig.course;
                    this.page.url = 'https://disqus_eb4a097382404ffe81300aac5744e3fe_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_eb4a097382404ffe81300aac5744e3fe_' + eBookConfig.course;
            disqus_url = 'http://disqus_eb4a097382404ffe81300aac5744e3fe_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_eb4a097382404ffe81300aac5744e3fe" onclick="disqus_eb4a097382404ffe81300aac5744e3fe_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_eb4a097382404ffe81300aac5744e3fe']").attr('data-disqus-identifier', 'disqus_eb4a097382404ffe81300aac5744e3fe_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.3" id="selection_ex_2" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_2 >
<div id=ex_6_2_question class="ac_question col-md-12">
<p>Give the <strong>logical opposites</strong> of these conditions.  You are not allowed to use the <code class="docutils literal notranslate"><span class="pre">not</span></code> operator.</p>
<ol class="arabic simple">
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;</span> <span class="pre">b</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">b</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">18</span>&#160; <span class="pre">and</span>&#160; <span class="pre">day</span> <span class="pre">==</span> <span class="pre">3</span></code></p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">a</span> <span class="pre">&gt;=</span> <span class="pre">18</span>&#160; <span class="pre">or</span>&#160; <span class="pre">day</span> <span class="pre">!=</span> <span class="pre">3</span></code></p></li>
</ol>

</div>
<textarea data-lang="python" id="ex_6_2_editor" data-question_label="7.10.4" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.5" id="selection_ex_3" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='q3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_3 >
<div id=ex_6_3_question class="ac_question col-md-12">
<p>Write a function which is given an exam mark, and it returns a string — the grade for that mark — according to this
scheme:</p>
<table class="docutils align-default">
<colgroup>
<col style="width: 58%" />
<col style="width: 42%" />
</colgroup>
<thead>
<tr class="row-odd"><th class="head"><p>Mark</p></th>
<th class="head"><p>Grade</p></th>
</tr>
</thead>
<tbody>
<tr class="row-even"><td><p>&gt;= 90</p></td>
<td><p>A</p></td>
</tr>
<tr class="row-odd"><td><p>[80-90)</p></td>
<td><p>B</p></td>
</tr>
<tr class="row-even"><td><p>[70-80)</p></td>
<td><p>C</p></td>
</tr>
<tr class="row-odd"><td><p>[60-70)</p></td>
<td><p>D</p></td>
</tr>
<tr class="row-even"><td><p>&lt; 60</p></td>
<td><p>F</p></td>
</tr>
</tbody>
</table>
<p>The square and round brackets denote closed and open intervals.
A closed interval includes the number, and open interval excludes it.   So 79.99999 gets grade C , but 80 gets grade B.</p>
<p>Test your function by printing the mark and the grade for a number of different marks.</p>

</div>
<textarea data-lang="python" id="ex_6_3_editor" data-question_label="7.10.6" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def getGrade(grade):
      #your code here


====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(getGrade(95),'A','Tested getGrade on input of 95')
        self.assertEqual(getGrade(85),'B','Tested getGrade on input of 85')
        self.assertEqual(getGrade(65),'D','Tested getGrade on input of 65')
        self.assertEqual(getGrade(79.99999),'C','Tested getGrade on input of 79.9999')
        self.assertEqual(getGrade(80),'B','Tested getGrade on input of 80')


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q3_question >
<div id=q3_question_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="q3_question_editor" data-question_label="7.10.7" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def grade(mark):
    if mark >= 90:
        return "A"
    else:
        if mark >= 80:
            return "B"
        else:
            if mark >= 70:
                return "C"
            else:
                if mark >= 60:
                    return "D"
                else:
                    return "F"

mark = 83
print( "Mark:", str(mark), "Grade:", grade(mark))
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_b9e6fd14629949e59da1a4ae827c0032_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_b9e6fd14629949e59da1a4ae827c0032_' + eBookConfig.course;
                    this.page.url = 'https://disqus_b9e6fd14629949e59da1a4ae827c0032_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_b9e6fd14629949e59da1a4ae827c0032_' + eBookConfig.course;
            disqus_url = 'http://disqus_b9e6fd14629949e59da1a4ae827c0032_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_b9e6fd14629949e59da1a4ae827c0032" onclick="disqus_b9e6fd14629949e59da1a4ae827c0032_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_b9e6fd14629949e59da1a4ae827c0032']").attr('data-disqus-identifier', 'disqus_b9e6fd14629949e59da1a4ae827c0032_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.8" id="selection_ex_4" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_4 >
<div id=ex_6_4_question class="ac_question col-md-12">
<p>Modify the turtle bar chart program from the previous chapter so that the bar for any value
of 200 or more is filled with red, values between [100 and 200) are filled yellow,
and bars representing values less than 100 are filled green.</p>

</div>
<textarea data-lang="python" id="ex_6_4_editor" data-question_label="7.10.9" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.10" id="selection_ex_5" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='q5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_5 >
<div id=ex_6_5_question class="ac_question col-md-12">
<p>In the turtle bar chart program, what do you expect to happen if one or more
of the data values in the list is negative?   Go back and try it out.  Change the
program so that when it prints the text value for the negative bars, it puts
the text above the base of the bar (on the 0 axis).</p>

</div>
<textarea data-lang="python" id="ex_6_5_editor" data-question_label="7.10.11" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer_ex_6_5 >
<div id=answer_ex_6_5_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer_ex_6_5_editor" data-question_label="7.10.12" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
import turtle

def drawBar(t, height):
    """ Get turtle t to draw one bar, of height. """
    t.begin_fill()               # start filling this shape
    if height < 0:
        t.write(str(height))
    t.left(90)
    t.forward(height)
    if height >= 0:
        t.write(str(height))
    t.right(90)
    t.forward(40)
    t.right(90)
    t.forward(height)
    t.left(90)
    t.end_fill()                 # stop filling this shape



xs = [48, -50, 200, 240, 160, 260, 220]  # here is the data
maxheight = max(xs)
minheight = min(xs)
numbars = len(xs)
border = 10

tess = turtle.Turtle()           # create tess and set some attributes
tess.color("blue")
tess.fillcolor("red")
tess.pensize(3)

wn = turtle.Screen()             # Set up the window and its attributes
wn.bgcolor("lightgreen")
if minheight > 0:
    lly = 0
else:
    lly = minheight - border

wn.setworldcoordinates(0-border, lly, 40*numbars+border, maxheight+border)


for a in xs:
    drawBar(tess, a)

wn.exitonclick()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_0118bd02de23462bafdb51beb4c85e44_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_0118bd02de23462bafdb51beb4c85e44_' + eBookConfig.course;
                    this.page.url = 'https://disqus_0118bd02de23462bafdb51beb4c85e44_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_0118bd02de23462bafdb51beb4c85e44_' + eBookConfig.course;
            disqus_url = 'http://disqus_0118bd02de23462bafdb51beb4c85e44_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_0118bd02de23462bafdb51beb4c85e44" onclick="disqus_0118bd02de23462bafdb51beb4c85e44_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_0118bd02de23462bafdb51beb4c85e44']").attr('data-disqus-identifier', 'disqus_0118bd02de23462bafdb51beb4c85e44_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.13" id="selection_ex_6" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_6 >
<div id=ex_6_6_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">findHypot</span></code>.  The function will be given the length of two sides of a right-angled triangle and it should return
the length of the hypotenuse.  (Hint:  <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">**</span> <span class="pre">0.5</span></code> will return the square root, or use <code class="docutils literal notranslate"><span class="pre">sqrt</span></code> from the math module)</p>

</div>
<textarea data-lang="python" id="ex_6_6_editor" data-question_label="7.10.14" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

def findHypot(a,b):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(findHypot(12.0,5.0),13.0,"Tested findHypot on inputs of 12.0 and 5.0")
        self.assertEqual(findHypot(14.0,48.0),50.0,"Tested findHypot on inputs of 14.0 and 48.0")
        self.assertEqual(findHypot(21.0,72.0),75.0,"Tested findHypot on inputs of 21.0 and 72.0")
        self.assertAlmostEqual(findHypot(1,1.73205),1.999999,2,"Tested findHypot on inputs of 1 and 1.73205")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.15" id="selection_ex_7" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='q7' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_7 >
<div id=ex_6_7_question class="ac_question col-md-12">
<p>Write a function called <code class="docutils literal notranslate"><span class="pre">is_even(n)</span></code> that takes an integer as an argument
and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the argument is an <strong>even number</strong> and <code class="docutils literal notranslate"><span class="pre">False</span></code> if
it is <strong>odd</strong>.</p>

</div>
<textarea data-lang="python" id="ex_6_7_editor" data-question_label="7.10.16" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

def is_even(n):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
     def testOne(self):
         self.assertEqual(is_even(10),True,"Tested is_even on input of 10")
         self.assertEqual(is_even(5),False,"Tested is_even on input of 5")
         self.assertEqual(is_even(1),False,"Tested is_even on input of 1")
         self.assertEqual(is_even(0),True,"Tested is_even on input of 0")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_answer >
<div id=q7_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="q7_answer_editor" data-question_label="7.10.17" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def is_even(n):
    if n % 2 == 0:
        return True
    else:
        return False

testEqual(is_even(10), True)
testEqual(is_even(5), False)
testEqual(is_even(1), False)
testEqual(is_even(0), True)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_7ae92646976d4953ac8a163f338b4358_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_7ae92646976d4953ac8a163f338b4358_' + eBookConfig.course;
                    this.page.url = 'https://disqus_7ae92646976d4953ac8a163f338b4358_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_7ae92646976d4953ac8a163f338b4358_' + eBookConfig.course;
            disqus_url = 'http://disqus_7ae92646976d4953ac8a163f338b4358_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_7ae92646976d4953ac8a163f338b4358" onclick="disqus_7ae92646976d4953ac8a163f338b4358_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_7ae92646976d4953ac8a163f338b4358']").attr('data-disqus-identifier', 'disqus_7ae92646976d4953ac8a163f338b4358_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.18" id="selection_ex_8" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_8 >
<div id=ex_6_8_question class="ac_question col-md-12">
<p>Now write the function <code class="docutils literal notranslate"><span class="pre">is_odd(n)</span></code> that returns <code class="docutils literal notranslate"><span class="pre">True</span></code> when <code class="docutils literal notranslate"><span class="pre">n</span></code> is odd
and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.</p>

</div>
<textarea data-lang="python" id="ex_6_8_editor" data-question_label="7.10.19" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

def is_odd(n):
    # your code here


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
     def testOne(self):
         self.assertEqual(is_odd(10),False,"Tested is_odd on input of 10")
         self.assertEqual(is_odd(5),True,"Tested is_odd on input of 5")
         self.assertEqual(is_odd(1),True,"Tested is_odd on input of 1")
         self.assertEqual(is_odd(0),False,"Tested is_odd on input of 0")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.20" id="selection_ex_9" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='q9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_9 >
<div id=ex_6_9_question class="ac_question col-md-12">
<p>Modify <code class="docutils literal notranslate"><span class="pre">is_odd</span></code> so that it uses a call to <code class="docutils literal notranslate"><span class="pre">is_even</span></code> to determine if its
argument is an odd integer.</p>

</div>
<textarea data-lang="python" id="ex_6_9_editor" data-question_label="7.10.21" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

def is_odd(n):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
     def testOne(self):
         self.assertEqual(is_odd(10),False,"Tested is_odd on input of 10")
         self.assertEqual(is_odd(5),True,"Tested is_odd on input of 5")
         self.assertEqual(is_odd(1),True,"Tested is_odd on input of 1")
         self.assertEqual(is_odd(0),False,"Tested is_odd on input of 0")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_answer >
<div id=q9_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="q9_answer_editor" data-question_label="7.10.22" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def is_even(n):
    if n % 2 == 0:
        return True
    else:
        return False

def is_odd(n):
    if is_even(n):
        return False
    else:
        return True

testEqual(is_odd(10), False)
testEqual(is_odd(5), True)
testEqual(is_odd(1), True)
testEqual(is_odd(0), False)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_9125792d0c264b65b4d9d72d25485ceb_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_9125792d0c264b65b4d9d72d25485ceb_' + eBookConfig.course;
                    this.page.url = 'https://disqus_9125792d0c264b65b4d9d72d25485ceb_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_9125792d0c264b65b4d9d72d25485ceb_' + eBookConfig.course;
            disqus_url = 'http://disqus_9125792d0c264b65b4d9d72d25485ceb_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_9125792d0c264b65b4d9d72d25485ceb" onclick="disqus_9125792d0c264b65b4d9d72d25485ceb_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_9125792d0c264b65b4d9d72d25485ceb']").attr('data-disqus-identifier', 'disqus_9125792d0c264b65b4d9d72d25485ceb_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.23" id="selection_ex_10" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_10 >
<div id=ex_6_10_question class="ac_question col-md-12">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">is_rightangled</span></code> which, given the length of three sides of a triangle,
will determine whether the triangle is right-angled.  Assume that the third argument to the
function is always the longest side.  It will return <code class="docutils literal notranslate"><span class="pre">True</span></code> if the triangle
is right-angled, or <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.</p>
<p>Hint: floating point arithmetic is not always exactly accurate,
so it is not safe to test floating point numbers for equality.
If a good programmer wants to know whether
<code class="docutils literal notranslate"><span class="pre">x</span></code> is equal or close enough to <code class="docutils literal notranslate"><span class="pre">y</span></code>, they would probably code it up as</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span>  <span class="nb">abs</span><span class="p">(</span><span class="n">x</span> <span class="o">-</span> <span class="n">y</span><span class="p">)</span> <span class="o">&lt;</span> <span class="mf">0.001</span><span class="p">:</span>      <span class="c1"># if x is approximately equal to y</span>
    <span class="o">...</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ex_6_10_editor" data-question_label="7.10.24" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def is_rightangled(a, b, c):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(is_rightangled(1.5,2.0,2.5),True,"Tested is_rightangled on inputs of 1.5, 2.0 and 2.5")
        self.assertEqual(is_rightangled(4.0,8.0,16.0),False,"Tested is_rightangled on inputs of 4.0, 8.0 and 16.0")
        self.assertEqual(is_rightangled(4.1,8.2,9.1678787077),True,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.1678787077")
        self.assertEqual(is_rightangled(4.1,8.2,9.16787),True,"Tested is_rightangled on inputs of 4.1, 8.2, and 9.16787")
        self.assertEqual(is_rightangled(4.1,8.2,9.168),False,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.168")
        self.assertEqual(is_rightangled(0.5,0.4,0.64031),True,"Tested is_rightangled on inputs of 0.5, 0.4 and 0.64031")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.25" id="selection_ex_11" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    <div id='q11' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_11 >
<div id=ex_6_11_question class="ac_question col-md-12">
<p>Extend the above program so that the sides can be given to the function in any order.</p>

</div>
<textarea data-lang="python" id="ex_6_11_editor" data-question_label="7.10.26" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

def is_rightangled(a, b, c):
    # your code here


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(is_rightangled(1.5,2.5,2.0),True,"Tested is_rightangled on inputs of 1.5, 2.5 and 2.0")
        self.assertEqual(is_rightangled(16.0,4.0,8.0),False,"Tested is_rightangled on inputs of 16.0, 4.0 and 8.0")
        self.assertEqual(is_rightangled(4.1,8.2,9.1678787077),True,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.1678787077")
        self.assertEqual(is_rightangled(4.1,9.16787,8.2),True,"Tested is_rightangled on inputs of 4.1, 9.16787 and 8.2")
        self.assertEqual(is_rightangled(4.1,8.2,9.168),False,"Tested is_rightangled on inputs of 4.1, 8.2 and 9.168")
        self.assertEqual(is_rightangled(0.5,0.64031,0.4),True,"Tested is_rightangled on inputs of 0.5, 0.64031 and 0.4")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_answer >
<div id=q11_answer_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="q11_answer_editor" data-question_label="7.10.27" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
from test import testEqual

def is_rightangled(a, b, c):
    is_rightangled = False

    if a > b and a > c:
        is_rightangled = abs(b**2 + c**2 - a**2) < 0.001
    elif b > a and b > c:
        is_rightangled = abs(a**2 + c**2 - b**2) < 0.001
    else:
        is_rightangled = abs(a**2 + b**2 - c**2) < 0.001
    return is_rightangled

testEqual(is_rightangled(1.5, 2.0, 2.5), True)
testEqual(is_rightangled(4.0, 8.0, 16.0), False)
testEqual(is_rightangled(4.1, 8.2, 9.1678787077), True)
testEqual(is_rightangled(4.1, 8.2, 9.16787), True)
testEqual(is_rightangled(4.1, 8.2, 9.168), False)
testEqual(is_rightangled(0.5, 0.4, 0.64031), True)
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_b25992fa70fc4e7581d84315df38d796_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_b25992fa70fc4e7581d84315df38d796_' + eBookConfig.course;
                    this.page.url = 'https://disqus_b25992fa70fc4e7581d84315df38d796_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_b25992fa70fc4e7581d84315df38d796_' + eBookConfig.course;
            disqus_url = 'http://disqus_b25992fa70fc4e7581d84315df38d796_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_b25992fa70fc4e7581d84315df38d796" onclick="disqus_b25992fa70fc4e7581d84315df38d796_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_b25992fa70fc4e7581d84315df38d796']").attr('data-disqus-identifier', 'disqus_b25992fa70fc4e7581d84315df38d796_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.28" id="selection_ex_12" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_12 >
<div id=ex_6_12_question class="ac_question col-md-12">
<p>3 criteria must be taken into account to identify leap years:</p>
<p>The year is evenly divisible by 4;</p>
<p>If the year can be evenly divided by 100, it is NOT a leap year, unless;</p>
<p>The year is also evenly divisible by 400. Then it is a leap year.</p>
<p>Write a function that takes a year as a parameter and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the year is a leap year, <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.</p>

</div>
<textarea data-lang="python" id="ex_6_12_editor" data-question_label="7.10.29" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true"    
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">
def isLeap(year):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(isLeap(1944),True,"Tested isLeap on an input of 1944")
        self.assertEqual(isLeap(2011),False,"Tested isLeap on an input of 2011")
        self.assertEqual(isLeap(1986),False,"Tested isLeap on an input of 1986")
        self.assertEqual(isLeap(1800),False,"Tested isLeap on an input of 1800")
        self.assertEqual(isLeap(1900),False,"Tested isLeap on an input of 1900")
        self.assertEqual(isLeap(2000),True,"Tested isLeap on an input of 2000")
        self.assertEqual(isLeap(2056),True,"Tested isLeap on an input of 2056")

myTests().main()
</textarea>
</div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" data-question_label="7.10.30" id="selection_ex_13#." >
    <ol start=13 class=arabic><li class="alert alert-warning">

    <div id='q13' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ex_6_13 >
<div id=ex_6_13_question class="ac_question col-md-12">
<p>Implement the calculator for the date of Easter.</p>
<p>The following algorithm computes the date for Easter Sunday for any year between 1900 to 2099.</p>
<p>Ask the user to enter a year.
Compute the following:</p>
<blockquote>
<div><ol class="arabic simple">
<li><p>a = year % 19</p></li>
<li><p>b = year % 4</p></li>
<li><p>c = year % 7</p></li>
<li><p>d = (19 * a + 24) % 30</p></li>
<li><p>e = (2 * b + 4 * c + 6 * d + 5) % 7</p></li>
<li><p>dateofeaster = 22 + d + e</p></li>
</ol>
</div></blockquote>
<p>Special note: The algorithm can give a date in April.  Also, if the year is one of four special
years (1954, 1981, 2049, or 2076) then subtract 7 from the date.</p>
<p>Your program should print an error message if the user provides a date that is out of range.</p>

</div>
<textarea data-lang="python" id="ex_6_13_editor" data-question_label="7.10.31" 
    data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true"   
    data-audio=''      
        data-gradebutton=true   data-wasm=/_static
    style="visibility: hidden;">

</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=answer_ex_6_13 >
<div id=answer_ex_6_13_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="answer_ex_6_13_editor" data-question_label="7.10.32" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
year = int(input("Please enter a year"))
if year >= 1900 and year <= 2099:
    a = year % 19
    b = year % 4
    c = year % 7
    d = (19*a + 24) % 30
    e = (2*b + 4*c + 6*d + 5) % 7
    dateofeaster = 22 + d + e

    if year == 1954 or year == 2981 or year == 2049 or year == 2076:
        dateofeaster = dateofeaster - 7

    if dateofeaster > 31:
        print("April", dateofeaster - 31)
    else:
        print("March", dateofeaster)
else:
    print("ERROR...year out of range")
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_2dfd6acf1ca849c2853dad606d1ba255_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course;
                    this.page.url = 'https://disqus_2dfd6acf1ca849c2853dad606d1ba255_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course;
            disqus_url = 'http://disqus_2dfd6acf1ca849c2853dad606d1ba255_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'https://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_2dfd6acf1ca849c2853dad606d1ba255" onclick="disqus_2dfd6acf1ca849c2853dad606d1ba255_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_2dfd6acf1ca849c2853dad606d1ba255']").attr('data-disqus-identifier', 'disqus_2dfd6acf1ca849c2853dad606d1ba255_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 7.9. Glossary' data-toggle="tooltip">
    <a href="Glossary.html" aria-label="7.9. Glossary">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 8. More About Iteration' data-toggle="tooltip" >
    <a href="../MoreAboutIteration/toctree.html" aria-label="8. More About Iteration">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2014 Brad Miller, David Ranum, Created using Runestone Interactive.
      Last updated on Nov 14, 2021.<br/>
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.5.5.
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>