<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>20.7. Testing with pytest &#8212; How to Think like a Computer Scientist: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/user-highlights.css?v=A2326E52" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.5.5" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/pytutor-embed.bundle.js?v=268B5C04"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=92D8BBF0"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.7/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="20.8. Glossary" href="Glossary.html" />
    <link rel="prev" title="20.6. Test-First Development" href="TestFirstDevelopment.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/sdmesa/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = false;
    eBookConfig.host = 'https://sdmesa.github.io/thinkcspy/' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/sdmesa';
    eBookConfig.course = 'thinkcspy';
    eBookConfig.basecourse = 'thinkcspy';
    eBookConfig.isLoggedIn = false;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.5.5';
  eBookConfig.jobehost = '';
  eBookConfig.proxyuri_runs = '';
  eBookConfig.proxyuri_files = '';
  eBookConfig.enable_chatcodes = false ? false : false;


</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
          <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="https://www.sdmesa.edu">&nbsp; </a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          How to Think like a Computer Scientist
        </a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- WWW external links dropdown -->

        <!--
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social" class="visuallyhidden">Social</span></i>
          </a>
        </li>-->

        <!-- end WWW external links dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/sdmesa/default/reportabug?course=thinkcspy&page=TestingWithpytest'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">20.6. </span>Test-First Development" data-toggle="tooltip">
    <a href="TestFirstDevelopment.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">20.8. </span>Glossary' data-toggle="tooltip" >              
    <a href="Glossary.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->







    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">
<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="testing-with-pytest">
<h1><span class="section-number">20.7. </span>Testing with pytest<a class="headerlink" href="#testing-with-pytest" title="Permalink to this headline">¶</a></h1>
<p>Writing automated unit tests is very helpful in reducing the effort needed to build software. However, the simple
approach described so far is inadequate to help programmers realize the full benefits of unit testing. In this section,
we introduce a unit test framework which addresses some practical issues that crop up when you try to apply
unit testing techniques in software development projects. Here are some of the issues with using plain assertion unit
tests:</p>
<ul class="simple">
<li><p>Simple assert-based tests don’t give very good diagnostic information when they fail. It would help to have better
reporting. For example, when an assert fails, it would help us in diagnosing the error to see what value the function
actually produced. An AssertionError doesn’t give us that information.</p></li>
<li><p>We need a better way to organize our unit test code. So far, I’ve suggested creating separate programs to hold
the functions under test together with their unit test code. But that isn’t practical for most projects. For example,
functions often need to call other functions in the program in order to do their work. It’s not convenient to
bring those auxiliary functions over into separate test programs.</p></li>
<li><p>We need a way to keep the unit test around and use it even after the function is first developed
to help us catch bugs that are inadvertently introduced into the function when we make modifications to it.</p></li>
</ul>
<p>Unit testing frameworks help to address these issues, by improving error reporting, providing a structure
for programmers to organize their unit tests, and making it possible to leverage existing unit tests when
making enhancements to functions. <strong>pytest</strong> is one unit testing framework that provides these benefits.</p>
<p>For our purposes, the attractive thing about pytest is that writing unit tests with pytest feels a lot like writing unit
tests using plain assert statements. The only difference is that you put your assert statements into test functions.
Here’s an example of how it works:</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_round6_pytest1 >
<div id=ac_round6_pytest1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_round6_pytest1_editor" data-question_label="20.7.1" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
def round6(num: float) -> int:
    """This function has a bug in it"""
    return int(num + .6)

# ---- automated unit test ----

def test_round6():
    assert round6(9.7) == 10
    assert round6(8.5) == 8

=====

from unittest.gui import TestCaseGui

class pyTests(TestCaseGui):

    def testOne(self):
        for item in globals():
            if item.startswith("test_"):
                try:
                    globals()[item]()
                    self.assertEqual(item + " passed", item + " passed", item + " passed")
                except Exception as e:
                    self.assertEqual(item + " failed", item + " passed", item + " failed: " + str(e))

pyTests().main()
</textarea>
</div>
</div>
<p>This code example defines two functions: the function to be tested, <code class="docutils literal notranslate"><span class="pre">round6</span></code>, and a function named <code class="docutils literal notranslate"><span class="pre">test_round6</span></code>
that contains unit test code. When using the pytest approach, you write your unit test as a function whose name must
start with the prefix <code class="docutils literal notranslate"><span class="pre">test_</span></code>. Inside the function, you write normal assert statements to test the desired function.
Notice that you do not write a line to <em>call</em> the unit test function. Instead, when you launch pytest to run the unit
tests, pytest scans your script and executes only the functions with the prefix <code class="docutils literal notranslate"><span class="pre">test_</span></code>.</p>
<p>This ActiveCode environment simulates pytest by scanning for and executing functions with a <code class="docutils literal notranslate"><span class="pre">test_</span></code> prefix when you
click <strong>Run</strong>. Go ahead and try it - rename the test_round6 function to <code class="docutils literal notranslate"><span class="pre">test_itworks</span></code> and try running the test again.</p>
<div class="section" id="organizing-pytest-functions">
<h2><span class="section-number">20.7.1. </span>Organizing pytest Functions<a class="headerlink" href="#organizing-pytest-functions" title="Permalink to this headline">¶</a></h2>
<p>The example above uses a single pytest function, with both asserts in the same pytest function. The
disadvantage of that approach is that the first failing assert prevents the rest of the asserts from
being tested.</p>
<p>If you want, you can write multiple pytest functions to test a single function. That way, when an
assert fails in one test function, the rest of the pytest functions can still run and report success or
failure.</p>
<p>You can name your pytest functions with names that indicate what they are testing. For
example, try changing the ActiveCode example above so that it defines two test functions: one named
<code class="docutils literal notranslate"><span class="pre">test_round6_rounds_up</span></code>, containing the first assert, and one named <code class="docutils literal notranslate"><span class="pre">test_round6_rounds_down</span></code>,
containing the second assert. Your code should look like this:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">test_round6_rounds_up</span><span class="p">():</span>
    <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">9.7</span><span class="p">)</span> <span class="o">==</span> <span class="mi">10</span>

<span class="k">def</span> <span class="nf">test_round6_rounds_down</span><span class="p">():</span>
    <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span> <span class="o">==</span> <span class="mi">8</span>
</pre></div>
</div>
<p>If you use good pytest function names, when a pytest function has an assertion failure, you can
easily tell what the problem was.</p>
</div>
<div class="section" id="using-pytest">
<h2><span class="section-number">20.7.2. </span>Using pytest<a class="headerlink" href="#using-pytest" title="Permalink to this headline">¶</a></h2>
<p>To use pytest, you must first install it using the <strong>pip</strong> command. Open your computer’s command line window
(not the Python interpreter) and enter the following command to install:</p>
<ul>
<li><p>Windows:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">pip</span> <span class="n">install</span> <span class="n">pytest</span>
</pre></div>
</div>
</li>
<li><p>Mac/Linux:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">pip3</span> <span class="n">install</span> <span class="n">pytest</span>
</pre></div>
</div>
</li>
</ul>
<p>After you have installed pytest, you run pytest unit tests from the command line window. To run pytest unit
tests, try copying the code from the ActiveCode example above and pasting it into a Python file named (ex.)
<strong>myround.py</strong>. Then, use the <strong>pytest</strong> command to run your tests by opening a command window,
navigating to the folder where you stored myround.py, and executing the following command:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">pytest</span> <span class="n">myround</span><span class="o">.</span><span class="n">py</span>
</pre></div>
</div>
</div>
<div class="section" id="understanding-pytest-failure-reports">
<h2><span class="section-number">20.7.3. </span>Understanding pytest Failure Reports<a class="headerlink" href="#understanding-pytest-failure-reports" title="Permalink to this headline">¶</a></h2>
<p>When you run the pytest command and an assertion fails, you see a report like this:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">===============================</span> <span class="n">FAILURES</span> <span class="o">================================</span>
<span class="n">______________________________</span> <span class="n">test_round6</span> <span class="n">______________________________</span>
    <span class="k">def</span> <span class="nf">test_round6</span><span class="p">():</span>
        <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">9.7</span><span class="p">)</span> <span class="o">==</span> <span class="mi">10</span>
<span class="o">&gt;</span>       <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span> <span class="o">==</span> <span class="mi">8</span>
<span class="n">E</span>       <span class="k">assert</span> <span class="mi">9</span> <span class="o">==</span> <span class="mi">8</span>
<span class="n">E</span>        <span class="o">+</span>  <span class="n">where</span> <span class="mi">9</span> <span class="o">=</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span>

<span class="n">myround</span><span class="o">.</span><span class="n">py</span><span class="p">:</span><span class="mi">8</span><span class="p">:</span> <span class="ne">AssertionError</span>
</pre></div>
</div>
<p>Let’s take a closer look at this report to understand what it’s telling you.</p>
<ol class="arabic">
<li><p>First, notice the line with the <code class="docutils literal notranslate"><span class="pre">&gt;</span></code> symbol:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">&gt;</span>       <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span> <span class="o">==</span> <span class="mi">8</span>
</pre></div>
</div>
<p>The <code class="docutils literal notranslate"><span class="pre">&gt;</span></code> symbol points to the line with the assertion that failed.</p>
</li>
<li><p>Next, notice the lines marked <code class="docutils literal notranslate"><span class="pre">E</span></code>:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">E</span>       <span class="k">assert</span> <span class="mi">9</span> <span class="o">==</span> <span class="mi">8</span>
<span class="n">E</span>        <span class="o">+</span>  <span class="n">where</span> <span class="mi">9</span> <span class="o">=</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span>
</pre></div>
</div>
<p>This indicates that the call to <code class="docutils literal notranslate"><span class="pre">round6(8.5)</span></code> returned the value 9, instead of the value 8.
The value <code class="docutils literal notranslate"><span class="pre">9</span></code> is the actual result of the function. Knowing the value actually produced by
the function can help you to troubleshoot the bug and correct the problem.</p>
</li>
</ol>
</div>
<div class="section" id="integrated-unit-testing-with-pytest">
<h2><span class="section-number">20.7.4. </span>Integrated Unit Testing with pytest<a class="headerlink" href="#integrated-unit-testing-with-pytest" title="Permalink to this headline">¶</a></h2>
<p>When you use the pytest framework, you can include pytest test functions in your main program, along with the rest of
your program code. This allows you to keep your tests together with the functions that they test, and you can run either
your program (using the python command) or the unit tests (using the pytest command).</p>
<p>Take a look at this example that shows a function (<code class="docutils literal notranslate"><span class="pre">round6</span></code>, containing a bug), together with a
unit test function (<code class="docutils literal notranslate"><span class="pre">test_round6</span></code>), and a main program that uses <code class="docutils literal notranslate"><span class="pre">round6</span></code>:</p>
<div class="highlight-python notranslate"><table class="highlighttable"><tr><td class="linenos"><div class="linenodiv"><pre> 1
 2
 3
 4
 5
 6
 7
 8
 9
10
11
12
13
14</pre></div></td><td class="code"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">round6</span><span class="p">(</span><span class="n">num</span><span class="p">:</span> <span class="nb">float</span><span class="p">)</span> <span class="o">-&gt;</span> <span class="nb">int</span><span class="p">:</span>
    <span class="k">return</span> <span class="nb">int</span><span class="p">(</span><span class="n">num</span> <span class="o">+</span> <span class="o">.</span><span class="mi">6</span><span class="p">)</span>

<span class="c1"># ---- automated unit test ----</span>

<span class="k">def</span> <span class="nf">test_round6</span><span class="p">():</span>
    <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">9.7</span><span class="p">)</span> <span class="o">==</span> <span class="mi">10</span>
    <span class="k">assert</span> <span class="n">round6</span><span class="p">(</span><span class="mf">8.5</span><span class="p">)</span> <span class="o">==</span> <span class="mi">8</span>

<span class="c1"># ----- main program follows -----</span>

<span class="k">if</span> <span class="vm">__name__</span> <span class="o">==</span> <span class="s1">&#39;__main__&#39;</span><span class="p">:</span>
    <span class="n">num</span> <span class="o">=</span> <span class="nb">float</span><span class="p">(</span><span class="nb">input</span><span class="p">(</span><span class="s1">&#39;Enter a value:&#39;</span><span class="p">))</span>
    <span class="nb">print</span><span class="p">(</span><span class="s1">&#39;The value rounded is: &#39;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(</span><span class="n">round6</span><span class="p">(</span><span class="n">num</span><span class="p">)))</span>
</pre></div>
</td></tr></table></div>
<p>Notice how the main program is inside the <code class="docutils literal notranslate"><span class="pre">if</span></code> statement on line 12. This if condition is true when the program is run
using the <strong>python</strong> command, and allows the main program to execute. When the unit tests are executed using the
<strong>pytest</strong> command, any top-level code outside a function in the python file gets executed when <strong>pytest</strong> scans the
script looking for unit test functions with a <code class="docutils literal notranslate"><span class="pre">test_</span></code> prefix. The <code class="docutils literal notranslate"><span class="pre">if</span></code> condition is false in this scenario, and that
prevents the main program from executing when <strong>pytest</strong> is scanning the script. If that explanation didn’t make total
sense, just remember: in order for pytest to work correctly, any code that is part of the main program must be inside an
<code class="docutils literal notranslate"><span class="pre">if</span></code> statement like the one in this example, so that it doesn’t interfere with pytests’s unit testing process.</p>
<p><strong>Check your understanding</strong></p>
<div id='tab_grade_unittest' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_grade_pytest >
<div id=ac_grade_pytest_question class="ac_question col-md-12">
<p>Write a pytest unit test function named <code class="docutils literal notranslate"><span class="pre">test_grade</span></code> to test a function
with the following specification. Your asserts should
check that the function produces an appropriate value
for each of the three postcondition cases.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">grade</span><span class="p">(</span><span class="n">score</span><span class="p">):</span>
    <span class="sd">&quot;&quot;&quot;Determines letter grade given a numeric score</span>

<span class="sd">    Precondition: 0 &lt;= `score` &lt;= 100</span>
<span class="sd">    Postcondition: Returns &#39;A&#39; if 90 &lt;= `score` &lt;= 100,</span>
<span class="sd">      &#39;B&#39; if 80 &lt;= `score` &lt; 90, &#39;F&#39; if 0 &lt;= `score` &lt; 80</span>
<span class="sd">    &quot;&quot;&quot;</span>
</pre></div>
</div>

</div>
<textarea data-lang="python" id="ac_grade_pytest_editor" data-question_label="20.7.4.1" 
     data-include="ac_grade_pytest_aux" data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
    style="visibility: hidden;">
# Write a pytest unit test function named ``test_grade``


====
from unittest.gui import TestCaseGui

testA = False
testB = False
testF = False
illegal = False

def grade(score):
    global illegal, testA, testB, testF

    if score > 100 or score < 0:
        illegal = True
        return ''
    elif score >= 90:
        testA = True
        return 'A'
    elif score >= 80:
        testB = True
        return 'B'
    else:
        testF = True
        return 'F'

class myTests(TestCaseGui):

    def testOne(self):
        code = self.getEditorText().replace(' ','').replace('"', '').replace("'", '')
        self.assertEqual(test_grade(), None, 'test_grade function defined' )
        self.assertTrue(testA and '==A' in code, "Assert tested 90..100")
        self.assertTrue(testB and '==B' in code, "Assert tested 80..90")
        self.assertTrue(testF and '==F' in code, "Assert tested 0..80")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
<p>The following is a suggested pytest unit test.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">test_grade</span><span class="p">():</span>
    <span class="k">assert</span> <span class="n">grade</span><span class="p">(</span><span class="mi">92</span><span class="p">)</span> <span class="o">==</span> <span class="s1">&#39;A&#39;</span>
    <span class="k">assert</span> <span class="n">grade</span><span class="p">(</span><span class="mi">85</span><span class="p">)</span> <span class="o">==</span> <span class="s1">&#39;B&#39;</span>
    <span class="k">assert</span> <span class="n">grade</span><span class="p">(</span><span class="mi">69</span><span class="p">)</span> <span class="o">==</span> <span class="s1">&#39;F&#39;</span>
</pre></div>
</div>
</div>
    </div>

</div>
</div>


  <div id=scprogresscontainer>
    You have attempted <span id=scprogresstotal></span> of <span id=scprogressposs></span> activities on this page <div id="subchapterprogress"></div>
  </div>
  
      
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - 20.6. Test-First Development' data-toggle="tooltip">
    <a href="TestFirstDevelopment.html" aria-label="20.6. Test-First Development">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - 20.8. Glossary' data-toggle="tooltip" >
    <a href="Glossary.html" aria-label="20.8. Glossary">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2014 Brad Miller, David Ranum, and Dave Parillo.
      Last updated on Nov 14, 2021.<br/>
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.5.5.
    </p>
  </div>
</footer>




<script> getSwitch(); </script>


  </body>
</html>